---
title: 'Chapter 2 – Identifying  Functional Issues'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Lesson 2.1 - Introduction to Functional Issues in Code Reviews

```yaml
type: VideoExercise
key: 876db58bbd
xp: 50
```

`@projector_key`
68da53846fe21cd14eaa9fbe784742cb
